---
title       : Lab 4b
description : Population and Sample Means


--- type:NormalExercise lang:r xp:100 skills:1 key:57cc190216
## Review random sampling

In this lab we'll be using another built-in R dataset called `ChickWeight` which contains data from an experiment on the effect of diet on chicks. 

- There are 4 variables in the data set: `weight`, `Time`, `Chick`, and `Diet`
- There are 578 subjects
- `ChickWeight` is already in your workspace


In this exercise we are going take random samples from `ChickWeight` using the `sample()` function.

Remember, to take a random sample and save it to an object, use the format:

`sample_name <- sample(dataset$variable, number)`.

Note: Remember that after we create an object we need to type the object name to view it *(I have already typed this for you in the R Script)*.

*** =instructions
- Take a random sample of 25 from the variable `weight` in the `ChickWeight` dataset, assign this to `sample_weight`.
- Take a random sample of 50 from the variable `Diet` in the `ChickWeight` dataset, assign this to `sample_diet`.
*** =hint
- For the 1st instruction, follow the format in the lesson but use ChickWeight$weight and 25

- For the 2nd instruction, follow the format in the lesson but use ChickWeight$Diet and 50
*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}
# take a random sample of 25 from the variable weight in ChickWeight dataset
sample_weight <- 



# Print out the values in sample_weight

sample_weight


# take a random sample of 50 from the variable Diet in ChickWeight dataset
sample_diet <- 



# Print out the values in sample_diet

sample_diet



```

*** =solution
```{r}
# take a random sample of 25 from the variable weight in ChickWeight dataset
sample_weight <- sample(ChickWeight$weight, 25)

# Print out the values in sample_weight
sample_weight

# take a random sample of 50 from the variable Diet in ChickWeight dataset
sample_diet <- sample(ChickWeight$Diet, 50)


# Print out the values in sample_diet
sample_diet



```

*** =sct
```{r}
test_object("sample_weight", incorrect_msg = "Make sure you use the sample() function on the right side of '<-' with the specified parameters; in this case we're using the ChickWeight dataset, weight variable, and number=25.")

test_output_contains("sample_weight", incorrect_msg = "Do not alter the line of code that says 'sample_weight'. This command will print the results to the R output. You can click the 'reverse' button to get back to the original code.")



test_object("sample_diet", incorrect_msg = "Make sure you use the sample() function on the right side of '<-' with the specified parameters; in this case we're using the ChickWeight dataset, Diet variable, and number=50.")

test_output_contains("sample_diet", incorrect_msg = "Do not alter the line of code that says 'sample_diet'. This command will print the results to the R output. You can click the 'reverse' button to get back to the original code.")


```



--- type:NormalExercise lang:r xp:100 skills:1 key:7e24794db5
## Population distribution


For this exercise, you are going to find the *population distribution* of the quantitative variable `weight` in the `ChickWeight` data set.

We will treat the `ChickWeight` data set as the entire population of interest.


*** =instructions
- Type `library(mosaic)` to load the *mosaic* package into your workspace. (You need this package to use the `favstats()` function)
- Use the `favstats()` function to get summary statistics of `weight` from the `ChickWeight` data set.
- Use the `hist()` function to create a histogram of `weight`.
- Click the 'Submit Answer' button and look at the R output
*** =hint
- Remember, to find summary statistics use format `favstats(dataset$variable)`.
- To get a histogram, use `hist(dataset$variable)`.
*** =pre_exercise_code
```{r}
library(mosaic)
```

*** =sample_code
```{r}
# Load the mosaic package


# Find summary statistics of chick weights


# Create a histogram of chick weights

```

*** =solution
```{r}
# Load the mosaic package
library(mosaic)

# Find summary statistics of chick weights
favstats(ChickWeight$weight)

# Create a histogram of chick weights
hist(ChickWeight$weight)

```

*** =sct
```{r}
test_student_typed("library(mosaic)",  not_typed_msg =  "Type library(mosaic) to load the mosaic package.")



test_function("favstats", args = "x", incorrect_msg = "Remember, to find summary statistics use format: favstats(dataset$variable) with specified parameters.")
test_function("hist", args = "x", incorrect_msg = "To create a histogram, use the format: hist(dataset$variable) with the specified variables.")

success_msg("Nice! Now look at the histogram of chick weights. Can you tell which way it is skewed?")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:4face42bf0
## Sample mean and distribution

Most of the time we don't know what the *population mean* is, but we can **estimate** this value with a random sample. 

In this exercise, you will:

* create a random sample of 75 from the `weight` variable in the `ChickWeight` data set, and,

* compare the *sample mean* `weight` of chicks to the *population mean* you found in the last exercise.


*** =instructions
- type `library(mosaic)` to load the *mosaic* package into your workspace
- use `sample()` to take a sample of 75 from the `weight` variable, assign this to `my_sample`
- use the `favstats()` function to find the mean and standard deviation of `my_sample`
- use the `hist()` function on `my_sample`
- click 'Submit Answer' and look at the R output.
- How do these values compare to the *population mean*?
*** =hint
- To get a random sample, use the format `sample(dataset$variable, #)` with the specified dataset, variable and number.

- `my_sample` is just a vector of 75 chick weights so to obtain summary statistics, all you need to put in the `favstats()` function is the sample name `my_sample`.


*** =pre_exercise_code
```{r}
library(mosaic)

```

*** =sample_code
```{r}
# Load the mosaic package


# Use the format sample(dataset$variable, #) to get a sample of 75 from the `weight` variable
my_sample <- 

# Find the mean and standard deviation of your new sample with favstats() 



# Create a histogram of your new sample

```

*** =solution
```{r}
# Load the mosaic package
library(mosaic)

# Use the format sample(dataset$variable, #) to get a sample of 75 from the `weight` variable
my_sample <- sample(ChickWeight$weight, 75)

# Find the mean and standard deviation of your new sample with favstats() 
favstats(my_sample)


# Create a histogram of your new sample
hist(my_sample)

```

*** =sct
```{r}
test_student_typed("library(mosaic)",  not_typed_msg =  "Type library(mosaic) to load the mosaic package.")


test_object("my_sample", incorrect_msg = "Make sure you use the sample(dataset$variable, number) function on the right side of '<-' with the specified parameters; in this case we're using the ChickWeight dataset, weight variable, and number=75.")





test_function("sample", args = c("x", "size"),  incorrect_msg = "To create an object called my_sample, type sample(ChickWeight$weight, 75) on the right side of '<-'.")

test_function("favstats", args = "x", incorrect_msg = "Remember, my_sample is just a sample from the weight variable so to get summary statistics type favstats(my_sample).")

test_function("hist", args = "x", incorrect_msg = "Remember, my_sample is just a sample from the weight variable so to get a histogram type hist(my_sample).")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:d7a8f825e2
## Confidence interval for a population mean

Assume that the population mean `weight` of chicks is unknown so we obtained a random sample. 

Finding the sample mean is okay but adding a margin of error gives us much more insight.

In this exercise, we are going to use R as a calculator to find a 95% confidence interval for the population mean *weight*.

Below is the formula for a confidence interval of a population mean.

$\bar{x} \pm t\ast \frac{s}{\sqrt{n}}$

*** =instructions
Suppose I sampled 75 of the chicks from `ChickWeight` and found some sample statistics:

* mean weight is 116.2
* standard deviation is 67.5
* for n = 75 and alpha = 0.05, the $t\ast$ critical value is 1.99


Use R as a calculator to find a 95% confidence interval for the population mean *weight* based on this sample data (Use the formula in the lesson).
*** =hint
Use the formula in the lesson with $\bar{x}$ = 116.2, $s$ = 67.5, $t\ast$ = 1.99, and n = 75.

Find 'upper bound' and 'lower bound' seperately.
*** =pre_exercise_code
```{r}


```

*** =sample_code
```{r}
#Lower confidence bound

#Upper confidence bound


```

*** =solution
```{r}
#Lower confidence bound
116.2 - 1.99 * (67.5/sqrt(75))

#Upper confidence bound
116.2 + 1.99 * (67.5/sqrt(75))


```

*** =sct
```{r}
test_output_contains("100.6895", incorrect_msg = "Follow the order of operations from your high school algebra class 'PEMDAS' and remember, when you want to square something in R, use the sqrt() function. For the lower confidence bound, subtract the margin of error from the sample mean.")

test_output_contains("131.7105", incorrect_msg = "Follow the order of operations from your high school algebra class 'PEMDAS' and remember, when you want to square something in R, use the sqrt() function. For the upper confidence bound, add the margin of error to the sample mean.")


success_msg("Good job. You finished Lab 4b practice course.")

```
