---
title       : Lab 7
description : Insert the chapter description here



--- type:NormalExercise lang:r xp:100 skills:1 key:616c48f5cf
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
