---
title       : Lab 6
description : Insert the chapter description here


--- type:NormalExercise lang:r xp:100 skills:1 key:79a37231a2
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
