---
title       : Lab 8
description : Correlation and Simple Linear Regression



--- type:NormalExercise lang:r xp:100 skills:1 key:283c94164f
## Gun violence data

In this pre-lab, we'll be using a data set called `gun_violence_us` which contains gun ownership and mortality statistics in the United States in 2014.

- there are 3 variables in this data set: `state`, `ownership_rate`, `mortality_rate`
- `gun_violence_us` is already in your workspace

Let's begin by exploring the data set.
*** =instructions
- Use the `head()` function on `gun_violence_us` to view the first six observations.
- Use the `summary()` function on `gun_violence_us` to obtain summary statistics for each variable.
*** =hint
Use the format `head(dataset)` for the first instruction and `summary(dataset)` for the second instruction.
*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")

```

*** =sample_code
```{r}
# Use the head() function on gun_violence_us to view the first six observations


# Use the summary() function on gun_violence_us to obtain summary statistics for each variable


```

*** =solution
```{r}
# Use the head() function on gun_violence_us to view the first six observations
head(gun_violence_us)

# Use the summary() function on gun_violence_us to obtain summary statistics for each variable
summary(gun_violence_us)

```

*** =sct
```{r}
test_function("head", args = "x", incorrect_msg = "Make sure to use `gun_violence_us` instead of 'dataset' inside the `head()` function.")

test_function("summary", args = "object", incorrect_msg = "Follow the format `summary(dataset)` with the specified data set to obtain a summary of the entire data set.")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:b81533a6c6
## Scatter plot review

Suppose we want to assess if there is a relationship between `ownership_rate` and `mortality_rate`.

* `ownership_rate`: percentage of adults in each state who own a gun

* `mortality_rate`: number of deaths per 100,000 people 

Let's start by visualizing the relationship between these variables with the `plot()` function.

Remember, to create a scatterplot use the format:

`plot(x = dataset$var1, y = dataset$var2)`
*** =instructions
Use the `plot()` function to visualize the relationship between `ownership_rate` and `mortality_rate` in the `gun_violence_us` data set.

* `ownership_rate` should be on the x-axis
* `mortality_rate` should be on the y-axis
*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
```

*** =sample_code
```{r}
# Use the plot() function to visualize the relationship between ownership_rate and mortality_rate


```

*** =solution
```{r}
# Use the plot() function to visualize the relationship between ownership_rate and mortality_rate
plot(x = gun_violence_us$ownership_rate, y = gun_violence_us$mortality_rate)


```

*** =sct
```{r}
test_function("plot", args = c("x", "y"), incorrect_msg = "Remember to follow the format in the lesson with the specified data set (`gun_violence_us`) and variables.")

success_msg("Nice! What trend do you observe in the plot?")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:23b4d73093
## Measuring correlation with the cor() function

We can use the `cor()` function to estimate the correlation between two quantitative variables with the format:

`cor(x = dataset$var1, y = dataset$var2)`

*Note: the order of var1 and var2 does not matter for the `cor()` function.*

*** =instructions
Use the `cor()` function to estimate the correlation coefficient between `ownership_rate` and `mortality_rate` from the `gun_violence_us` data set.
*** =hint
Follow the `cor()` format in the lesson with the `gun_violence_us` data set and `ownership_rate`/`mortality_rate` in any order. 
*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
```

*** =sample_code
```{r}
# Estimate the correlation between ownership_rate and mortality_rate


```

*** =solution
```{r}
# Estimate the correlation between ownership_rate and mortality_rate
cor(x = gun_violence_us$ownership_rate, y = gun_violence_us$mortality_rate)


```

*** =sct
```{r}
test_function("cor", args = c("x", "y"), eval = F, incorrect_msg = "Make sure to follow the format: `cor(x = dataset$var1, y = dataset$var2)` with the correct data set and variables.")
```



--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:e07efd5860
## Quick check 1

Using the command:

`cor(x = gun_violence_us$ownership_rate, y = gun_violence_us$mortality_rate)` 

will produce the same results as the command: 

`cor(x = gun_violence_us$mortality_rate, y = gun_violence_us$ownership_rate)`.

*** =instructions
- FALSE
- TRUE
*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
msg_bad <- "That is not correct, the order of var1 and var2 does not matter for the `cor()` function."
msg_success <- "Exactly!"
test_mc(correct = 2, feedback_msgs = c(msg_bad, msg_success))

```
--- type:NormalExercise lang:r xp:100 skills:1 key:76ae5ee967
## Simple linear regression

We can learn more about the relationship between two quantitative variables by estimating a simple linear regression model with the `lm()` function in the format:

`model1 <- lm(dataset$y_var ~ dataset$x_var)`

`model1`

* `y_var` = the response variable
* `x_var` = the explanatory variable


We are saving the model results as an object called `model1`. Remember, we need to type the object name to view it.

*** =instructions
* Use the `lm()` function to estimate a linear regression model from the `gun_violence_us` data set and store these regression results in an object called `gun_model` where:

    - `y_var` = `mortality_rate`
    - `x_var` = `ownership_rate`

*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
```

*** =sample_code
```{r}
# Use the lm() function to calculate a linear model of ownership_rate and mortality_rate
gun_model <- 

# Type the name of the model (gun_model) to view the results



```

*** =solution
```{r}
# Use the lm() function to calculate a linear model of ownership_rate and mortality_rate
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

# Type the name of the model (gun_model) to view the results
gun_model

```

*** =sct
```{r}
test_object("gun_model", incorrect_msg = "Use the `lm()` function on the right side of the tilde with the correct data set (`gun_violence_us`), y.var  (`mortality_rate`) and x.var  (`ownership_rate`).")

test_student_typed("gun_model", not_typed_msg = "Do not remove or alter the line of code that prints the contents of `gun_model`.") 

success_msg("Good work. Now look at the R output and find the estimate for the intercept and slope.")
```






--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:9fc077c265
## Quick check 2

The regression model from the last exercise `gun_model` is stored in your workspace; type `gun_model` in your R script to view it.

What is the estimated regression equation?

*** =instructions
- $\hat{y} = 4.573 + 20.492x$
- $\hat{y} = 4.573 + 20.492$
- $\hat{y} = 20.492 + 4.573x$
- $\hat{y} = 20.492 + 20.492x$

*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

```

*** =sct
```{r}
msg_bad <- "That is not correct."
msg_success <- "Exactly!"
test_mc(correct = 1, feedback_msgs = c(msg_success, msg_bad, msg_bad, msg_bad))

```

--- type:NormalExercise lang:r xp:100 skills:1 key:3605a42869
## Testing if the population slope is non-zero

To view more comprehensive regression model results, produce a `summary` of the model.  For example, if we had executed

`model1 <- lm(dataset$y_var ~ dataset$x_var)`

then 

`summary(model1)` would produce more comprehensive regression results.


Suppose we want to test whether there is a linear association between `ownership_rate` and `mortality_rate` in the `gun_violence_us` data set.

Our hypotheses are:

$H\_0:$ $\beta\_{1} = 0$

$H\_a:$ $\beta\_{1} \neq 0$

where $\beta\_{1}$ = the population slope of the regression line.

These results can be obtained with the `summary` function.


*** =instructions
- *Note: `gun_model` is already in your workspace.*
- Summarize `gun_model` using the `summary()` function.
- Click the 'Submit Answer' button and look at the p-value of the test.
*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

```

*** =sample_code
```{r}
# Use the summary() function on gun_model


```

*** =solution
```{r}
# Use the summary() function on gun_model
summary(gun_model)


```

*** =sct
```{r}
test_function("summary", args = "object", incorrect_msg = "Type: `summary(gun_model)` below the first comment.")
success_msg("Good job.")
```



--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:d94f04f494
## Quick check 3

The results from the last exercise are stored in an object called `model_results`; in your R console type `model_results` to view these.

At the alpha = 0.05 significance level, how would we conclude this test?

Recall our hypotheses are:

$H\_0:$ $\beta\_{1} = 0$

$H\_a:$ $\beta\_{1} \neq 0$

where $\beta\_{1}$ = the slope of the regression

Based on the p-value of `_____`, `_____` the null hypothesis.  We `_____` have evidence that the population slope differs from zero.

*** =instructions
- 0.00022; reject; do
- 0.00022; fail to reject; do not
- 6.24e-08; reject; do
- 6.24e-08; fail to reject; do not

*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)
model_results <- summary(gun_model)
```

*** =sct
```{r}
msg_bad <- "That is not correct. The p-value is almost zero so we would reject- we can see that ownership rate and mortality rate have a positive linear association so the slope is greater than zero."
msg_success <- "Exactly!"
test_mc(correct = 3, feedback_msgs = c(msg_bad, msg_bad, msg_success, msg_bad))

```
--- type:NormalExercise lang:r xp:100 skills:1 key:c7ec6c9c9b
## Confidence intervals for regression coefficients

To obtain confidence intervals for the regression coefficients, we can use the `confint()` function with the format:

`confint(model_name)`

*** =instructions
Use the `confint()` function to calculate confidence intervals for the `gun_model` regression coefficients (intercept, slope).
*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

```

*** =sample_code
```{r}
# Use confint() to obtain a confidence interval for the gun_model regression coefficients



```

*** =solution
```{r}
# Use confint() to obtain a confidence interval for the gun_model regression coefficients
confint(gun_model)


```

*** =sct
```{r}
test_function("confint", args = "object", incorrect_msg = "Make sure to follow the format in the lesson with `gun_model` instead of `model_name`.")
```
--- type:NormalExercise lang:r xp:100 skills:1 key:cded916a25
## Add lines to a plot with abline()

Follow the subsequent steps to add an estimated regression equation line to a scatter plot:

First, estimate the regression equation:

`model_name <- lm(dataset$y_var ~ dataset$x_var)`
    
Second, produce the scatterplot:

`plot(x = dataset$x_var, y = dataset$y_var)`

Third, use the `abline()` function with the following format to add the estimated regression line to the plot:

`abline(model_name)`




*** =instructions
- The first line of code creates a scatter plot of `ownership_rate` and `mortality_rate` from the `gun_violence_us` data set. **Do not alter this code.**

- Use the `abline()` function to add the `gun_model` regression line to the plot.
*** =hint

*** =pre_exercise_code
```{r}
source("http://www.openintro.org/stat/data/gun_violence_us.R")

```

*** =sample_code
```{r}
# Code to estimate the regression equation.  Do not change.
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

# Code to create a scatter plot of ownership_rate and mortality_rate. Do not change.
plot(gun_violence_us$ownership_rate, gun_violence_us$mortality_rate)


# Add the gun_model regression line to the scatter plot


```

*** =solution
```{r}
# Code to estimate the regression equation.  Do not change.
gun_model <- lm(gun_violence_us$mortality_rate ~ gun_violence_us$ownership_rate)

# Code to create a scatter plot of ownership_rate and mortality_rate. Do not change.
plot(gun_violence_us$ownership_rate, gun_violence_us$mortality_rate)


# Add the gun_model regression line to the scatter plot
abline(gun_model)


```

*** =sct
```{r}
test_student_typed("plot(gun_violence_us$ownership_rate, gun_violence_us$mortality_rate)", not_typed_msg = "Do not remove or alter the line of code that plots `ownership_rate` and `mortality_rate`. Instead, do the second instruction.")

test_function("abline", args = "a", incorrect_msg = "Make sure to follow the format in the lesson with `gun_model` instead of `model_name`.")

```

--- type:PlainMultipleChoiceExercise lang:r xp:50 skills:1 key:ecd40fc973
## LAST EXERCISE OF LAB 8

**This is the last exercice of Lab 8.**

Correlation can be used to estimate the strength of the linear association between:

*** =instructions
- one quantitative variable and $\mu_{0}$
- two quantitative variables
- two categorical variables
- one quantitative and one categorical variable

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}
msg_bad <- "That is not correct. What kind of variables were `ownership_rate` and `mortality_rate`?"
msg_success <- "Exactly! You have completed Lab 8 practice course."
test_mc(correct = 2, feedback_msgs = c(msg_bad, msg_success, msg_bad, msg_bad))

```
