---
title       : Lab 2
description : Using R to Explore Data Sets

--- type:NormalExercise lang:r xp:100 skills:1 key:f8b383874f
## Getting Started with Data

In this exercise, you will learn how to identify fundamental attributes of a data set.

I've already uploaded the data set `arbuthnot` into your workspace so if you type  `arbuthnot` into your R console, the entire data will come up in the output.

Sometimes the data can be very long so instead of viewing the entire thing, we can get a jist of the data by typing `head(dataset)` to get the first 6 entries or `tail(dataset)` to get the last 6 entries. 

If we just want to know the variables in the data set, we can type `names(dataset)` into the console.

*** =instructions
- Use the `head()` function to get the first 6 entries of the data set
- Use `tail()` to get the last 6 entries.
- Find the variable names by using the `names()` function.

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
