---
title       : Lab 7
description : Inference for Categorical Data



--- type:NormalExercise lang:r xp:100 skills:1 key:616c48f5cf
## One sample z-test

In this course, we will use the one sample z-test to test whether a population *proportion* differs from a specific value.

To perform a one sample z-test, we'll be using a data set called `nc` which contains data on the North Carolina birth records in 2004. 

- there are 5 variables in this data set but we will use 
- there are 154 subjects 
- `nc` is already in your workspace

This data set contains a categorical variable called `lowbirthweight` with 2 levels:

- `low` (low birthweight)
- `not low` (not low birthweight)

Suppose that a previous study estimated the low birthweight to be 10%. 


*** =instructions
- Test whether low birthweight in this population differs from 10%.

*** =hint

*** =pre_exercise_code
```{r}
source("https://www.openintro.org/stat/data/nc.R")
```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
