---
title       : Lab 3
description : Insert the chapter description here

--- type:NormalExercise lang:r xp:100 skills:1 key:fe98ce0ad3
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
